---
title: Test
description: Test
---

## Sample exercise

```yaml
type: NormalExercise 
lang: python
xp: 100 
skills: 2
key: 148efc63c1   
```


changed kjbkj


`@instructions`
sdsadadasdasdsadasdasdas

`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


---

## mp3

```yaml
type: VideoExercise 
xp: 50 
key: a3edb91bde   
```

`@projector_key`
2b26a811dbd8a609868733488a9814ee
---

## Title

```yaml
type: NormalExercise 
xp: 100 
key: a3edb91bdf   
```


This is the new exercise in this chapter


`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: 05f48a0294   
```

`@projector_key`
c57a30f043f834531721cdbb0add4d6f
