---
title: '2 chapter'
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise
key: 7c5d3aad0b
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: NormalExercise
key: 863d630ab0
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: NormalExercise
key: 979ed172c8
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: NormalExercise
key: 51f80fb55a
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: NormalExercise
key: ba3025752f
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`


---

## dasda

```yaml
type: NormalExercise
key: 55defdae77
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`
