---
title: Test
description: Test
---

## changed

```yaml
type: NormalExercise
key: 148efc63c1
lang: python
xp: 100
skills: 2
```

hello

`@instructions`
sdsadadasdasdsadasdasdas

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## mp3

```yaml
type: VideoExercise
key: a3edb91bde
xp: 50
```

`@projector_key`
2b26a811dbd8a609868733488a9814ee

---

## Title

```yaml
type: NormalExercise
key: a3edb91bdf
xp: 100
```

This is the new exercise in this chapter

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 05f48a0294
xp: 50
```

`@projector_key`
c57a30f043f834531721cdbb0add4d6f
