---
title: Test
description: Test
---

## Sample exercise

```yaml
type: NormalExercise 
lang: python
xp: 100 
skills: 2
key: 148efc63c1   
```


changed kjbkj


`@instructions`
sdsadadasdas

`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: a3edb91bde   
```

`@projector_key`
2b26a811dbd8a609868733488a9814ee
