---
title: '2'
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 7c5d3aad0b   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 863d630ab0   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 979ed172c8   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 51f80fb55a   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: ba3025752f   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 55defdae77   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


