---
title: 'Insert the chapter title here'
description: 'Insert the chapter description here'
---

## Insert exercise title here

```yaml
type: VideoExercise
key: 35d1482992
xp: 50
```

`@projector_key`
5ea9a54cf213692af7ece631cbdbc087

---

## Insert exercise title here

```yaml
type: VideoExercise
key: ae17d2adbc
xp: 50
```

`@projector_key`
cb792e452c222adc1f18767cc16e69bd

---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: 6e9eec2a71
xp: 50
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 9c67129e7f
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: 6f762ba673
xp: 50
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: VideoExercise
key: 6eac6e6faa
xp: 50
```

`@projector_key`
1929604626e43caece781461673defa6

---

## Insert exercise title here

```yaml
type: NormalExercise
key: a867133e9b
xp: 100
```

last in chapter 6

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`
