---
title: Test
description: Test
---

## Sample exercise

```yaml
type: NormalExercise 
lang: python
xp: 100 
skills: 2
key: 148efc63c1   
```


a change from github another


`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`

