---
title: '5'
description: ""
---

## Insert exercise title here

```yaml
type: VideoExercise
key: 0ea38fed66
xp: 50
```

`@projector_key`
cab8b85dfeff54d3d40e045819bbe10a

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 4b2636edb7
xp: 100
```

first in chapter 5

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: d9e9d79007
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
