---
title: 'Insert the chapter title here'
description: 'Insert the chapter description here'
---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: 35d1482992   
```

`@projector_key`
5ea9a54cf213692af7ece631cbdbc087
---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: ae17d2adbc   
```

`@projector_key`
cb792e452c222adc1f18767cc16e69bd
---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise 
xp: 50 
key: 6f762ba673   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: 6eac6e6faa   
```

`@projector_key`
1929604626e43caece781461673defa6
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: a867133e9b   
```


last in chapter 6


`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


