---
  title: "2"
  description: ""
---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: 0ea38fed66   
```

`@projector_key`
cab8b85dfeff54d3d40e045819bbe10a
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: d9e9d79007   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```

