---
title: '4'
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 5ec50b45a3   
```


first in chapter 4


`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


`@possible_answers`


`@feedback`


