---
title: '4'
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise
key: 5ec50b45a3
xp: 100
```

first in chapter 4

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

`@possible_answers`


`@feedback`
